# Assignment 9
## Due June 27th at 12:00

Please note: 

- Read the instructions in the exercise PDF and in this notebook carefully.
- Add your solutions *only* at `YOUR CODE HERE`/`YOUR ANSWER HERE` and remove the corresponding `raise NotImplementedError()`.
- Do not chance the provided code and text, if not stated.
- Do not *add* or *delete* cells.
- Do not `import` additional functionality. 
- Before submitting: Please make sure, that your notebook can be executed from top to bottom `Menu -> Kernel -> Restart & Run all`. 

# Exercise 1 (Exact recovery from a PCA is impossible, 2+1 points)

## a)
\begin{align}
    W &\in \mathbb{R}^{d \times n} \quad d<n\\
    \implies& \text{Rank}(W) \leq d \text{(not full rank)}\\
    \implies& \text{nullspace}(W) \neq \emptyset \\
    \implies& \exists u,v \in \text{nullspace}(W) \\
    \text{ s.t. } u&\neq v \text{ but } Wu = 0 = Wv
\end{align}

## b) 
For any input vector $u$ there exists at least one different input vector $v$ which would result in the same decoded vector. Thus, perfect reconstruction is impossible.

# Exercise 2: PCA and Random Projections

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd

from numpy.testing import assert_equal, assert_almost_equal
from sklearn.metrics.pairwise import euclidean_distances
from sklearn.random_projection import GaussianRandomProjection
from sklearn.decomposition import PCA
```

```python
# load the USPS dataset
usps_df = pd.read_csv('USPS.csv', sep='\t',)
xs_usps = usps_df.drop(['Label'], axis=1).values
ys_usps = usps_df['Label'].values

xs_usps.shape, ys_usps.shape
```

Output:
```
((1000, 256), (1000,))
```

#### Here you can implement PCA and Random Projections yourself. This will be awarded with 1.5 bonus points per method. If you decide to do so, you have to use your implementation in this exercise. If you do not want to implement this, simply remove the NotImplementedError().

<!-- grade_id: cell-3df0929058fe4b42, points: 0, grade: true, solution: true -->
```python
# Your PCA and random projection implementation HERE (OPTIONAL)
# YOUR CODE HERE
```

### (a) Reduce the dataset's dimensionality to d=2

<!-- grade_id: cell-c3313227fd7771a9, points: 0, grade: false, solution: true -->
```python
xs_usps_pca = ... # an np.array of shape 1000x2

# YOUR CODE HERE

plt.scatter(*xs_usps_pca.T, c=ys_usps)
plt.title('PCA Projection of USPS data')
plt.show()
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-4-44e8b8c33baf> in <module>
      3 # YOUR CODE HERE
      4 
----> 5 plt.scatter(*xs_usps_pca.T, c=ys_usps)
      6 plt.title('PCA Projection of USPS data')
      7 plt.show()

AttributeError: 'ellipsis' object has no attribute 'T'
```

<!-- grade_id: cell-34107104eb4060fc, points: 1, grade: true, solution: true -->
```python
xs_usps_random = ... # an np.array of shape 1000x2

# YOUR CODE HERE

plt.scatter(*xs_usps_random.T, c=ys_usps)
plt.title('Random Projection of USPS data')
plt.show()
```

### (b) Visualize prinicipal components and random directions

<!-- grade_id: cell-1fdd232b87f205d8, points: 1, grade: true, solution: true -->
```python
fig, axis = plt.subplots(2, 2)

# YOUR CODE HERE

plt.show()
```

### (c) Visualize reconstructed images

<!-- grade_id: cell-671019a2a8c38fb8, points: 0, grade: false, solution: true -->
```python
def reconstruct(xs_original, xs_projected):
    ''' Returns: A np.array of the same shape as xs_original
    '''

    dim = xs_projected.shape[1]
    R = np.zeros((256, dim))
    for i in range(256):
        R[i,:] = np.linalg.lstsq(xs_projected, xs_original[:,i], rcond=None)[0] 
    return (R @ xs_projected.T).T
```

<!-- grade_id: cell-df46991e5565e095, points: 3, grade: true, solution: false -->
```python
# assert correctness of reconstruct
np.random.seed(42)
W = np.random.normal(size=(256,2))/np.sqrt(2)
xs_reconstructed = reconstruct(xs_usps, xs_usps @ W)
assert_almost_equal(np.linalg.norm(xs_reconstructed), 28102.566390212647)

random_projection = GaussianRandomProjection(n_components=256)
xs_usps_random = random_projection.fit_transform(xs_usps)
assert (xs_usps - reconstruct(xs_usps, xs_usps_random)).max() < 1e-8
```

```python
# we want to reconstruct digit number 100 - visualize it
digit = xs_usps[100,:]
plt.imshow(digit.reshape(16, 16))
plt.show()
```

<!-- grade_id: cell-61fd1cb7a74f910e, points: 1, grade: true, solution: true -->
```python
dimensions = [1, 10, 50, 100, 256]
fig, axis = plt.subplots(2, 5, figsize=(15,5))

for idx, dim in enumerate(dimensions):
    # YOUR CODE HERE
    
    # axis[0, idx].imshow( pca reconstruction of image 100 for the given dimension )
    # axis[1, idx].imshow( random projection reconstruction of image 100 for the given dimension )
    axis[0, idx].set_title(f'PCA, d={dim}')
    axis[0, idx].axis('off')
    axis[1, idx].set_title(f'Random, d={dim}')
    axis[1, idx].axis('off')

plt.show()
```

### (d) Reconstruction Error of PCA and Random Projections

<!-- grade_id: cell-22874c8512852033, points: 1, grade: true, solution: true -->
```python
dimensions = [5, 25, 50, 100, 150, 200, 256]

errors_pca = ... 
errors_random = ...

# YOUR CODE HERE

plt.plot(dimensions, errors_random, 'o-', label='Random Projection')
plt.plot(dimensions, errors_pca, 'x--', label='PCA')
plt.xlabel('Embedding Dimension')
plt.xticks(dimensions)
plt.ylabel('Average Reconstruction Error')
plt.legend(loc='best')
plt.show()
```

### (e) Average and maximum distortion of distances

<!-- grade_id: cell-cd7be51626143865, points: 0, grade: false, solution: true -->
```python
def pairwise_distortions(xs_orig, xs_new):
    # YOUR CODE HERE

def average_distortion(xs_orig, xs_new):
    return pairwise_distortions(xs_orig, xs_new).mean()
    
def maximum_distortion(xs_orig, xs_new):
    return pairwise_distortions(xs_orig, xs_new).max()
```

<!-- grade_id: cell-f01fe96a39416594, points: 1, grade: true, solution: false -->
```python
xs_orig = np.arange(10).reshape(-1, 1)
xs_new = xs_orig * xs_orig

# checking correctness of distortion functions
assert_almost_equal(average_distortion(xs_orig, xs_orig), 1.)
assert_almost_equal(maximum_distortion(xs_new, xs_new), 1.)
assert_almost_equal(average_distortion(xs_orig, xs_new), 9.)
assert_almost_equal(maximum_distortion(xs_orig, xs_new), 17.)
p_usps = pairwise_distortions(xs_usps, PCA(2).fit_transform(xs_usps))
assert_equal(p_usps.shape, (499500, ))
```

<!-- grade_id: cell-f9a0b7a30c5633bd, points: 1, grade: true, solution: true -->
```python
dimensions = [5, 10, 25,  50]

avg_distortion_pca = ...
avg_distortion_random = ...

# YOUR CODE HERE

plt.plot(dimensions, avg_distortion_random, 'o-', label='Random Projection')
plt.plot(dimensions, avg_distortion_pca, 'x--', label='PCA')
plt.xlabel('Embedding Dimension')
plt.xticks(dimensions)
plt.ylabel('Average Distortion')
plt.legend(loc='best')
plt.show()

max_distortion_pca = ...
max_distortion_random = ...

# YOUR CODE HERE

plt.plot(dimensions, max_distortion_random, 'o-', label='Random Projection')
plt.plot(dimensions, max_distortion_pca, 'x--', label='PCA')
plt.xlabel('Embedding Dimension')
plt.xticks(dimensions)
plt.ylabel('Maximum Distortion')
plt.legend(loc='best')
plt.show()

```

### (f) Multiply the first dimension of the data by 500 and repeat (d) and (e).

```python
errors_pca = ... 
errors_random = ...

#YOUR CODE HERE

plt.plot(dimensions, errors_random, 'o-', label='Random Projection')
plt.plot(dimensions, errors_pca, 'x--', label='PCA')
plt.xlabel('Embedding Dimension')
plt.xticks(dimensions)
plt.ylabel('Average Reconstruction Error')
plt.legend(loc='best')
plt.show()
```

```python
avg_distortion_pca = ...
avg_distortion_random = ...

# YOUR CODE HERE

plt.plot(dimensions, avg_distortion_random, 'o-', label='Random Projection')
plt.plot(dimensions, avg_distortion_pca, 'x--', label='PCA')
plt.xlabel('Embedding Dimension')
plt.xticks(dimensions)
plt.ylabel('Average Distortion')
plt.legend(loc='best')
plt.show()

max_distortion_pca = ...
max_distortion_random = ...

# YOUR CODE HERE

plt.plot(dimensions, max_distortion_random, 'o-', label='Random Projection')
plt.plot(dimensions, max_distortion_pca, 'x--', label='PCA')
plt.xlabel('Embedding Dimension')
plt.xticks(dimensions)
plt.ylabel('Maximum Distortion')
plt.legend(loc='best')
plt.show()
```

### (g) Summarize the advantages and disadvantages of PCA and Random Projections

<!-- grade_id: cell-4e6f5405d2380844, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Exercise 3: Kernel PCA

```python
from sklearn.gaussian_process.kernels import RBF

def get_circle_data(n=50, r=1):
    
    np.random.seed(n)
    
    thetas = np.random.uniform(0, 2*np.pi, n)
    data = np.array([r * np.sin(thetas), r * np.cos(thetas)]).T + np.random.normal(scale=1/8, size=(n,2))
    
    return data
```

```python
data = np.concatenate([get_circle_data(100, 2),
                       get_circle_data(100, 4),
                       get_circle_data(100, 6)])

labels = np.concatenate([np.zeros(100),
                         np.ones(100),
                         2 * np.ones(100)])
```

```python
fig = plt.figure()
fig.set_size_inches((5,5))
color = ['orange' if l == 1 else 'red' if l == 2 else 'b' for l in labels]
plt.scatter(data[:, 0], data[:, 1], c=color)
plt.show()
```

## a)

```python
def kernel_PCA(xs, kernel, l=2):
    
   # YOUR CODE HERE
```

## b)

```python
# apply kernel PCA and 2d plot

# YOUR CODE HERE
```

```python
# 3d plot

# YOUR CODE HERE
```